# CSCI 3155: Assignment 3 

Topics covered: ASTs, pattern matching, higher order functions.

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

__Your Name Here__

## Problem 1 (30 points)

### 1A (10 points) 

This question asks you to implement the following grammar for a very simple programming language in scala.  The programming language allows you to declare variables and assign them values.

Example Program:

~~~
x1 := 10 # create a variable x1 and assign 10 to the variable x1.
x2 := 20 + x1 * x1 # create a variable x2 and assign 20 + x1 * x1
x3 := x1 + x2 # create a variable x3 and assign x1 + x2
x4 := x3 # create a variable x4 and assign x3 to it

(x1 + x2 + x3) # Final Return Expression that is the value of the whole program.
~~~

Please use constructors with the same name as the nonterminals in the grammar.



$$\newcommand\nt[1]{\textbf{#1}}$$
$$ \begin{array}{rcl}
\nt{Expr} & \Rightarrow & Plus(\nt{Expr}, \nt{Expr})\\
&|&  Minus(\nt{Expr}, \nt{Expr})\\
&|& Star(\nt{Expr}, \nt{Expr})\\
&|& Var(\nt{String}) \\
&|& Const(\nt{Double}) \\[10pt]
\nt{Statement} & \Rightarrow & Assignment(\nt{String}, \nt{Expr})\\[5pt]
\nt{Program} & \Rightarrow &  CalcProgram( \nt{Statement}* ,\  \nt{Expr})\\[5pt]
\end{array}$$


Complete the scala definition for the grammar above. Please use scala datatypes for Double and String.
Also, please use List(..) whenever Kleene star is used.


<!-- grade_id: p1a-ans, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Star(e1: Expr, e2: Expr) extends Expr
case class Var(s: String) extends Expr
case class Const(d: Double) extends Expr

//BEGIN SOLUTION
sealed trait Statement
case class Assignment(s: String, e: Expr) extends Statement

sealed trait Program
case class CalcProgram(lst: List[Statement], e: Expr) extends Program
//END SOLUTION

```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mStar[39m
defined [32mclass[39m [36mVar[39m
defined [32mclass[39m [36mConst[39m
defined [32mtrait[39m [36mStatement[39m
defined [32mclass[39m [36mAssignment[39m
defined [32mtrait[39m [36mProgram[39m
defined [32mclass[39m [36mCalcProgram[39m
```

<!-- grade_id: p1a-test1, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val v1 = Assignment("x", Const(2.0))
val v2 = Assignment("y", Plus(Var("x"), Const(4.0)))
val v3 = Assignment("z", Var("y"))
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mv1[39m: [32mAssignment[39m = [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m))
[36mv2[39m: [32mAssignment[39m = [33mAssignment[39m([32m"y"[39m, [33mPlus[39m([33mVar[39m([32m"x"[39m), [33mConst[39m([32m4.0[39m)))
[36mv3[39m: [32mAssignment[39m = [33mAssignment[39m([32m"z"[39m, [33mVar[39m([32m"y"[39m))
```

<!-- grade_id: p1a-t2, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val v1 = Assignment("x", Const(2.0))
val v2 = Assignment("y", Plus(Var("x"), Const(4.0)))
val v3 = Assignment("z", Var("y"))
val p1 = CalcProgram(List(v1, v2, v3), Var("z"))
val p2 = CalcProgram(List(), Var("y"))

passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mv1[39m: [32mAssignment[39m = [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m))
[36mv2[39m: [32mAssignment[39m = [33mAssignment[39m([32m"y"[39m, [33mPlus[39m([33mVar[39m([32m"x"[39m), [33mConst[39m([32m4.0[39m)))
[36mv3[39m: [32mAssignment[39m = [33mAssignment[39m([32m"z"[39m, [33mVar[39m([32m"y"[39m))
[36mp1[39m: [32mCalcProgram[39m = [33mCalcProgram[39m(
  [33mList[39m(
    [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m)),
    [33mAssignment[39m([32m"y"[39m, [33mPlus[39m([33mVar[39m([32m"x"[39m), [33mConst[39m([32m4.0[39m))),
    [33mAssignment[39m([32m"z"[39m, [33mVar[39m([32m"y"[39m))
  ),
  [33mVar[39m([32m"z"[39m)
)
[36mp2[39m: [32mCalcProgram[39m = [33mCalcProgram[39m([33mList[39m(), [33mVar[39m([32m"y"[39m))
```

### Part 1B (10 points)

In class we saw how to evaluate an expression given a map from identifiers to their values. 

~~~
def evalExpr(e: Expr, env: Map[String, Double]): Double = ...
~~~

In this assignment,  you are asked to handle assignment statement. The idea here is that the
assignment statement takes an environment and modifies it.

~~~
def evalStatement(s: Statement, env: Map[String, Double]): Map[String, Double] = ...
~~~

As an example consider the map  env

~~~
Map("x" -> 20, "y" -> 35)
~~~

The assignment statement z := x + y

~~~
Assignment("z", Plus(Var("x"), Var("y")))
~~~

Should return a map 

~~~
Map("x"-> 20, "y" -> 35, "z" -> 55)
~~~

The strategy should be as follows: (a) evaluate the expression in the RHS of the assignment statement, 
(b) add the entry that maps the declared variable to the new value and (c) return this map.

https://docs.scala-lang.org/overviews/collections-2.13/maps.html



<!-- grade_id: p1b-ans, points: 0, grade: false, solution: true -->
```scala
def evalExpr(e: Expr, env: Map[String, Double]): Double = {
    def binFun(e1: Expr, e2: Expr, op: (Double, Double) => Double): Double = {
        val v1 = evalExpr(e1, env)
        val v2 = evalExpr(e2, env)
        op(v1, v2)
    }
    
    e match {
        case Plus(e1, e2) => binFun(e1, e2, _ + _ )
        case Minus(e1, e2) => binFun(e1, e2, _ - _ )
        case Star(e1, e2) => binFun(e1, e2, _ * _ )
        case Var(v) => {
            if (env.contains(v)){
                env(v)
            } else {
                throw new IllegalArgumentException(s"$v undefined variable.")
            }
        }
        case Const(f) => f
    }
}

def evalStatement(s: Statement, env: Map[String, Double]): Map[String, Double] = s match {
    case Assignment(varName, rhsExpr) => {
        //BEGIN SOLUTION
        val d = evalExpr(rhsExpr, env)
        env + ( varName -> d)
        //END SOLUTION
    }
}

```

Output:
```
defined [32mfunction[39m [36mevalExpr[39m
defined [32mfunction[39m [36mevalStatement[39m
```

<!-- grade_id: p1b-test, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
def testMap(m1: Map[String, Double], varname: String, testval: Double) = {
    m1.contains(varname) && m1(varname) == testval
}

val m1 = Map[String, Double]()
val v1 = Assignment("x", Const(2.0))
val m2 = evalStatement(v1, m1)
assert(testMap(m2, "x", 2.0), "TEST 1 PASSED!")

val v2 = Assignment("y", Plus(Var("x"), Const(4.0)))
val m3 = evalStatement(v2, m2)
assert(testMap(m3, "x", 2.0), "TEST 2.1 PASSED!")
assert(testMap(m3, "y", 6.0), "TEST 2.2 PASSED!")


val v3 = Assignment("z", Star(Var("y"), Var("x")))
val m4 = evalStatement(v3, m3)
assert(testMap(m4, "x", 2.0), "TEST 3.1 PASSED!")
assert(testMap(m4, "y", 6.0), "TEST 3.2 PASSED!")
assert(testMap(m4, "z", 12.0), "TEST 3.3 PASSED!")

passed(10)

//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
defined [32mfunction[39m [36mtestMap[39m
[36mm1[39m: [32mMap[39m[[32mString[39m, [32mDouble[39m] = [33mMap[39m()
[36mv1[39m: [32mAssignment[39m = [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m))
[36mm2[39m: [32mMap[39m[[32mString[39m, [32mDouble[39m] = [33mMap[39m([32m"x"[39m -> [32m2.0[39m)
[36mv2[39m: [32mAssignment[39m = [33mAssignment[39m([32m"y"[39m, [33mPlus[39m([33mVar[39m([32m"x"[39m), [33mConst[39m([32m4.0[39m)))
[36mm3[39m: [32mMap[39m[[32mString[39m, [32mDouble[39m] = [33mMap[39m([32m"x"[39m -> [32m2.0[39m, [32m"y"[39m -> [32m6.0[39m)
[36mv3[39m: [32mAssignment[39m = [33mAssignment[39m([32m"z"[39m, [33mStar[39m([33mVar[39m([32m"y"[39m), [33mVar[39m([32m"x"[39m)))
[36mm4[39m: [32mMap[39m[[32mString[39m, [32mDouble[39m] = [33mMap[39m([32m"x"[39m -> [32m2.0[39m, [32m"y"[39m -> [32m6.0[39m, [32m"z"[39m -> [32m12.0[39m)
```

### Part 1C (10 points)

Write an evaluator for a program (list of statements + return expression) that evaluates the program and returns the final value of the return expression after the statements in the program are run one by one. We explain how to evaluate a program step by step on an example program

~~~
x := 15 //Stmt1
y := x - 10 // Stmt 2
z := y + 5 // Stmt 3
x+ y+ z // Return expression
~~~

This program should return the value 30. The key idea is this:

a) Start with an empty environment map to begin with.
b) Execute each statement in the list of statements from beginning to end.
   When you execute the statement, do so with the current map and the return value is the next map.
c) Evaluate the return expression on the final map after you have executed all the statements.

Complete the code below. You can use loops and var for now. We will show you how to avoid it later.

Here is how you iterate through each element of a list in scala: 

~~~
for (item <- lst){
    // Do stuff with item
}
~~~

<!-- grade_id: p1c-ans, points: 0, grade: false, solution: true -->
```scala
def evalProgram(p: Program): Double = p match {
    case CalcProgram(lstOfStatement, returnExpr) => {
        //BEGIN SOLUTION
        var env = Map[String,Double]()
        for (stmt <- lstOfStatement){
            env = evalStatement(stmt, env)
        }
        evalExpr(returnExpr, env)
        //END SOLUTION
    }
}
```

Output:
```
defined [32mfunction[39m [36mevalProgram[39m
```

<!-- grade_id: p1c-test1, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val v1 = Assignment("x", Const(2.0))
val v2 = Assignment("y", Plus(Var("x"), Const(4.0)))
val v3 = Assignment("z", Var("y"))
val p1 = CalcProgram(List(v1, v2, v3), Var("z"))
print(evalProgram(p1))
assert(evalProgram(p1) == 6)
passed(5)
//END TEST
```

Output:
```
6.0
*** Tests Passed (5 points) ***
[36mv1[39m: [32mAssignment[39m = [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m))
[36mv2[39m: [32mAssignment[39m = [33mAssignment[39m([32m"y"[39m, [33mPlus[39m([33mVar[39m([32m"x"[39m), [33mConst[39m([32m4.0[39m)))
[36mv3[39m: [32mAssignment[39m = [33mAssignment[39m([32m"z"[39m, [33mVar[39m([32m"y"[39m))
[36mp1[39m: [32mCalcProgram[39m = [33mCalcProgram[39m(
  [33mList[39m(
    [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m)),
    [33mAssignment[39m([32m"y"[39m, [33mPlus[39m([33mVar[39m([32m"x"[39m), [33mConst[39m([32m4.0[39m))),
    [33mAssignment[39m([32m"z"[39m, [33mVar[39m([32m"y"[39m))
  ),
  [33mVar[39m([32m"z"[39m)
)
```

<!-- grade_id: p1c-test2, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val vv1 = Assignment("x", Const(2.0))
val vv2 = Assignment("y", Star(Var("x"), Var("x")))
val vv3 = Assignment("z", Star(Var("y"), Var("y")))
val vv4 = Assignment("w", Star(Var("z"), Var("z")))
val p2 = CalcProgram(List(vv1, vv2, vv3, vv4), Var("w"))
print(evalProgram(p2))
assert(evalProgram(p2) == 256)
passed(5)
//END TEST
```

Output:
```
256.0
*** Tests Passed (5 points) ***
[36mvv1[39m: [32mAssignment[39m = [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m))
[36mvv2[39m: [32mAssignment[39m = [33mAssignment[39m([32m"y"[39m, [33mStar[39m([33mVar[39m([32m"x"[39m), [33mVar[39m([32m"x"[39m)))
[36mvv3[39m: [32mAssignment[39m = [33mAssignment[39m([32m"z"[39m, [33mStar[39m([33mVar[39m([32m"y"[39m), [33mVar[39m([32m"y"[39m)))
[36mvv4[39m: [32mAssignment[39m = [33mAssignment[39m([32m"w"[39m, [33mStar[39m([33mVar[39m([32m"z"[39m), [33mVar[39m([32m"z"[39m)))
[36mp2[39m: [32mCalcProgram[39m = [33mCalcProgram[39m(
  [33mList[39m(
    [33mAssignment[39m([32m"x"[39m, [33mConst[39m([32m2.0[39m)),
    [33mAssignment[39m([32m"y"[39m, [33mStar[39m([33mVar[39m([32m"x"[39m), [33mVar[39m([32m"x"[39m))),
    [33mAssignment[39m([32m"z"[39m, [33mStar[39m([33mVar[39m([32m"y"[39m), [33mVar[39m([32m"y"[39m))),
    [33mAssignment
```

## Problem 2 (20 points)

### A (5 points)
We defined lists in the class. Write a recursive procedure to get the nth element of the list or throw an `IllegalArgumentException` if $n < 0$ or $n >= \text{length of list} $. Assume $n=0$ obtains the very first element
and $n = \text{length of list} -1$ yields very last element. 

<!-- grade_id: p2a-ans, points: 0, grade: false, solution: true -->
```scala
sealed trait NumList
case object Nil extends NumList
case class Cons(n: Int, l: NumList) extends NumList 

def getNthElement(lst: NumList, n: Int): Int = {
    //BEGIN SOLUTION
    lst match {
        case Nil => throw new IllegalArgumentException()
        case Cons(elt, rest) => if (n == 0){ return elt } 
                                else if (n < 0){ throw new IllegalArgumentException ()}
                                else { getNthElement(rest, n-1) }
    }
    //END SOLUTION
}
```

Output:
```
defined [32mtrait[39m [36mNumList[39m
defined [32mobject[39m [36mNil[39m
defined [32mclass[39m [36mCons[39m
defined [32mfunction[39m [36mgetNthElement[39m
```

<!-- grade_id: p2a-test, points: 5, grade: true, solution: false -->
```scala
val l1 = Nil
val l2 = Cons(1, Cons(-1, Nil))
val l3 = Cons(1, Cons(2, l2))
val l4 = Cons(0, Cons(4, Cons(8, l3)))

val test1 = try {
    getNthElement(Nil, 3);
    assert(false, "Test 1 : getNthElement(Nil, 3) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
} 

assert(getNthElement(l2, 0) == 1, "Test2: getNthElement(l2, 0)  failed (expected answer = 1)")
assert(getNthElement(l3, 3) == -1, "Test3: getNthElement(l3, 3)  failed (expected answer = -1)")
assert(getNthElement(l4, 2) == 8, "Test4: getNthElement(l4, 2)  failed (expected answer = 8)")

val test2 = try {
    getNthElement(l4, 8);
    assert(false, "Test 5 : getNthElement(l4, 8) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
}

passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36ml1[39m: [32mNil[39m = Nil
[36ml2[39m: [32mCons[39m = [33mCons[39m([32m1[39m, [33mCons[39m([32m-1[39m, Nil))
[36ml3[39m: [32mCons[39m = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m1[39m, [33mCons[39m([32m-1[39m, Nil))))
[36ml4[39m: [32mCons[39m = [33mCons[39m([32m0[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m8[39m, [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m1[39m, [33mCons[39m([32m-1[39m, Nil)))))))
[36mtest1[39m: [32mAny[39m = [32m"OK"[39m
[36mtest2[39m: [32mAny[39m = [32m"OK"[39m
```

### B (7 points)
Write a recursive procedure that returns true if the list has the Fibonacci property. I.e, every element at position $i \geq 2$ is the sum of the two preceding elements. Note that the property is trivially true for lists of sizes $0$ and $1$.

<!-- grade_id: p2b-ans, points: 0, grade: false, solution: true -->
```scala
def isFibonacciList(lst: NumList): Boolean = {
    //BEGIN SOLUTION
    lst match {
        case Nil => true
        case Cons(i1, Nil) => true
        case Cons(i1, Cons(i2, Nil)) => true
        case Cons(i1, (rest @ Cons(i2, Cons(i3, _))) ) if (i3 == i1 + i2) => isFibonacciList(rest)
        case _ => false
    }
    //END SOLUTION
}
```

Output:
```
defined [32mfunction[39m [36misFibonacciList[39m
```

<!-- grade_id: p2b-test, points: 7, grade: true, solution: false -->
```scala
val l1 = Cons(12, Cons(25, Cons(37, Nil)))
assert(isFibonacciList(l1), 
       "Test case 1 :  isFibonacciList(l1) -- should return true")

val l2 = Cons(14, Cons(-1, Cons(13, l1 )))
assert(isFibonacciList(l2), 
       "Test case 2 :  isFibonacciList(l2) -- should return true")

val l3 = Cons(7, Cons(7, l2))

assert(!isFibonacciList(l3), 
       "Test case 3 :  isFibonacciList(l3) -- should return false")

val l4 = Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Nil))))))
assert(isFibonacciList(l2), 
       "Test case 4:  isFibonacciList(l4) -- should return true")

passed(7)
```

Output:
```
*** Tests Passed (7 points) ***
[36ml1[39m: [32mCons[39m = [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m37[39m, Nil)))
[36ml2[39m: [32mCons[39m = [33mCons[39m([32m14[39m, [33mCons[39m([32m-1[39m, [33mCons[39m([32m13[39m, [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m37[39m, Nil))))))
[36ml3[39m: [32mCons[39m = [33mCons[39m(
  [32m7[39m,
  [33mCons[39m([32m7[39m, [33mCons[39m([32m14[39m, [33mCons[39m([32m-1[39m, [33mCons[39m([32m13[39m, [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m37[39m, Nil)))))))
)
[36ml4[39m: [32mCons[39m = [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, Nil))))))
```

### C (8 points)
Write a recursive function `filterNumList(l: NumList, f: Int => Boolean): NumList` that takes in a `NumList` and a function `f: Int => Boolean`. 

1. It should return a new list that consist of all elements of the list `l` that return `true` when the function `f` is called on them.
2. The returned list elements must preserve the same order as in the original list.

<!-- grade_id: p2c-ans, points: 0, grade: false, solution: true -->
```scala
// BEGIN SOLUTION
def filterNumList(l: NumList, f: Int => Boolean): NumList = l match {
    case Nil => Nil
    case Cons(j, rest) if f(j) => Cons(j, filterNumList(rest, f)) 
    case Cons(_, rest) => filterNumList(rest, f)
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mfilterNumList[39m
```

<!-- grade_id: p2c-test1, points: 8, grade: true, solution: false -->
```scala
val l1 = Cons(12, Cons(25, Cons(37, Nil)))
def f1(j: Int): Boolean =  j <= 25 && j >= 12
assert(filterNumList(l1, f1) == Cons(12, Cons(25, Nil)), "Test 1 failed.")


val l2 = Cons(22, Cons(135, Cons(137, l1)))
def f2(j: Int): Boolean =  j % 5 == 0
assert(filterNumList(l2, f2) == Cons(135, Cons(25, Nil)), "Test 2 failed.")

def f3(j: Int): Boolean =  j >= 210
assert(filterNumList(l2, f3) == Nil, "Test 3 failed.")
assert(filterNumList(Nil, f3) == Nil, "Test 4 failed.")

val l4 = Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Nil))))))

def f4(j: Int): Boolean =  j <= 0
assert(filterNumList(l4, f4) == l4, "Test 5 failed")

passed(8)
```

Output:
```
*** Tests Passed (8 points) ***
[36ml1[39m: [32mCons[39m = [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m37[39m, Nil)))
defined [32mfunction[39m [36mf1[39m
[36ml2[39m: [32mCons[39m = [33mCons[39m([32m22[39m, [33mCons[39m([32m135[39m, [33mCons[39m([32m137[39m, [33mCons[39m([32m12[39m, [33mCons[39m([32m25[39m, [33mCons[39m([32m37[39m, Nil))))))
defined [32mfunction[39m [36mf2[39m
defined [32mfunction[39m [36mf3[39m
[36ml4[39m: [32mCons[39m = [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, [33mCons[39m([32m0[39m, Nil))))))
defined [32mfunction[39m [36mf4[39m
```