```scala

```

# CSCI 3155: Assignment 9

Topics: 
- Type Checking
- Basics of Objects
- Generics

Readings: Notes posted on moodle.  Chapters from Odersky

- Chapter 4 (revise)
- Chapter 10
- Chapter 11 (take note of Any, AnyVal, Null, Nothing)
- Chapter 12 
- Chapter 19

```scala
////////////////////////////////////////////////////////////////////////////////////
///////////////////////////////////////////////////////////////////////////////////
//__Name__:  Jason Nguyen - Assignment 9 - Sriram Sankaranarayanan - During CoVid-19 Outbreak 2020//
////////////////////////////////////////////////////////////////////////////////////
///////////////////////////////////////////////////////////////////////////////////
```

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (25 points)

In this problem, we will extend our type system for handling references.  Here is a stripped down version of Lettuce that involves references.

$$\begin{array}{rcl}
\mathbf{Expr} & \rightarrow & Const(\mathbf{Double}) \\
& | & Ident(\mathbf{Identifier})\\
& | & Plus(\mathbf{Expr}, \mathbf{Expr})\\
& | & Div(\mathbf{Expr}, \mathbf{Expr}) \\
& | & Geq(\mathbf{Expr}, \mathbf{Expr}) \\
& | & And(\mathbf{Expr}, \mathbf{Expr}) \\
& | & IfThenElse(\mathbf{Expr}, \mathbf{Expr}, \mathbf{Expr})\\
& | & Let(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}, \mathbf{Expr})\\
& | & FunDef(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}) \\
& | & FunCall(\mathbf{Expr}, \mathbf{Expr}) \\
& | & \color{blue}{NewRef(\mathbf{Expr})} \\
& | & \color{blue}{DeRef(\mathbf{Expr})} \\
& | & \color{blue}{AssignRef(\mathbf{Expr}, \mathbf{Expr})} \\
\end{array}$$

We have a type system 
$$\begin{array}{rcl}
\mathbf{Type} & \rightarrow & NumType & \leftarrow \text{written as } \ num \\
& | & BoolType & \leftarrow \text{written as } \ bool \\
& | & FunType(\mathbf{Type} , \mathbf{Type} ) & \leftarrow \text{written as } \ t_1 \Rightarrow t_2 \\
& | & \color{red}{RefType(\mathbf{Type})} & \leftarrow \text{written as } \ ref(t_1) \\
\end{array}$$


### A (5 points)
Copy the program below to the next cell and add the missing type annotations labeled type1, type2, and type3 to the the program.


~~~
let x: type1 = 25 in 
    let y : type2 = NewRef( - x ) in 
     let f : type3 = NewRef( function (z: ref(num)) DeRef(z) - x >= 45 ) in 
          let z : ref(bool) = DeRef(f)(y)
~~~

<!-- grade_id: cell-128df73da0a393cd, points: 5, grade: true, solution: true -->
type1: NumType type2: RefType(NumType) type2: RefType(FunType(RefType(NumType), BoolType))

### B (10 points): Let us write some inference rules for working with ref types.
$$\newcommand\typeOf{\mathbf{typeOf}}$$
$$\newcommand\semRule[3]{\begin{array}{c} #1 \\ \hline #2 \\ \end{array}\;(\text{#3}) }$$

Recall from notes on "Types and Type Checking" that $\typeOf(\texttt{e}, \alpha)$ is the type of an 
expression $\texttt{e}$ under type environment $\alpha$. The type environment maps identifiers in the current scope to  their annotated types.


Let us write a rule for NewRef.

$$\semRule{ \typeOf(\texttt{e}, \alpha) = t,\; t \not= \mathbf{typeerror} }{\typeOf(\texttt{NewRef(e)}, \alpha) = ref(t) }{newref-ok}$$

It says that if $\texttt{e}$ receives type $t$ under type environment $\alpha$ and it is not a type error, then $\texttt{NewRef(e)}$ must receive the type $ref(t)$ under $\alpha$.


(i) Complete the missing terms for the rule for `DeRef` OK rule.
$$\semRule{\typeOf(\texttt{e}, \alpha) = ref(t)}{\typeOf(\texttt{DeRef(e)}, \alpha) = \color{red}{???_1}}{deref-ok}$$

(ii) Complete the missing terms for the rule for `DeRef` error.
$$\semRule{\typeOf(\texttt{e}, \alpha) = t,\ t \not= \color{red}{???_2} }{\typeOf(\texttt{DeRef(e)}, \alpha) = \mathbf{typeerror}}{deref-nok}$$

(iii) Complete the missing terms for `AssignRef` OK rule.
$$\semRule{\typeOf(\texttt{e1}, \alpha) = \color{red}{???_3}, \typeOf(\texttt{e2}, \alpha) = t}{\typeOf(\texttt{AssignRef(e1, e2)}, \alpha) =   \color{red}{???_4}}{assignref-ok}$$


Write your answers in the next cell.

<!-- grade_id: cell-dd5541e0eaef1e6f, points: 10, grade: true, solution: true -->
Term 1: t
Term 2: ref(t)
Term 3: ref(t)
Term 4: t

### C (10 points): Write the type checker with references by filling in the missing code.

<!-- grade_id: cell-a9e5e31b52e7ee1b, points: 0, grade: false, solution: true -->
```scala
sealed trait Type
case object NumType extends Type
case object BoolType extends Type
case class FunType(t1: Type, t2: Type) extends Type
// TODO: Write a new case class for RefType
//??? // YOUR CODE HERE
case class RefType(t1: Type) extends Type        //Grammar

sealed trait Program
sealed trait Expr

case class Const(f: Double) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Geq(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(x: String, xType: Type, e1: Expr, e2: Expr) extends Expr
case class FunDef(id: String, idType: Type, e: Expr) extends Expr
case class FunCall(calledFun: Expr, argExpr: Expr) extends Expr
case class NewRef(e: Expr) extends Expr
case class DeRef(e: Expr) extends Expr
case class AssignRef(e1: Expr, e2: Expr) extends Expr
case class TopLevel(e: Expr) extends Program

def typeEquals(t1: Type, t2: Type): Boolean = t1 == t2
case class TypeErrorException(s: String) extends Exception

```

Output:
```
defined [32mtrait[39m [36mType[39m
defined [32mobject[39m [36mNumType[39m
defined [32mobject[39m [36mBoolType[39m
defined [32mclass[39m [36mFunType[39m
defined [32mclass[39m [36mRefType[39m
defined [32mtrait[39m [36mProgram[39m
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mIfThenElse[39m
defined [32mclass[39m [36mLet[39m
defined [32mclass[39m [36mFunDef[39m
defined [32mclass[39m [36mFunCall[39m
defined [32mclass[39m [36mNewRef[39m
defined [32mclass[39m [36mDeRef[39m
defined [32mclass[39m [36mAssignRef[39m
defined [32mclass[39m [36mTopLevel[39m
defined [32mfunction[39m [36mtypeEquals[39m
defined [32mclass[39m [36mTypeErrorException[39m
```

<!-- grade_id: cell-809dce336587273e, points: 0, grade: false, solution: true -->
```scala

def typeOf(e: Expr, alpha: Map[String, Type]): Type = {
    def checkType(opName: String, e1: Expr, t1: Type, e2: Expr, t2: Type, resType: Type): Type = {
        val t1hat = typeOf(e1, alpha)
        if (! typeEquals(t1hat, t1)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t1, obtained $t1hat")
        }
        
        val t2hat = typeOf(e2, alpha)
        if (! typeEquals(t2hat, t2)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t2, obtained $t2hat")
        }
        
        resType
    }
    
    e match {
        case Const(f) => NumType
        case Ident(s) => {if (alpha contains s)
                             alpha(s)
                          else 
                             throw TypeErrorException(s"Unknown identifier $s")}
        case Plus(e1, e2) =>  checkType("Plus", e1,  NumType, e2, NumType, NumType)
        case Minus(e1, e2) => checkType("Minus",e1,  NumType, e2, NumType, NumType)
        case Geq(e1, e2) => checkType("Geq", e1,  NumType, e2, NumType, BoolType)
        case IfThenElse(e, e1, e2) => {
            val t = typeOf(e, alpha)
            if (t == BoolType){
                val t1 = typeOf(e1, alpha)
                val t2 = typeOf(e2, alpha)
                if (typeEquals(t1, t2))
                    t1
                else 
                    throw TypeErrorException(s"If then else returns unequal types $t1 and $t2")
            } else {
                throw TypeErrorException(s"If then else condition expression not boolean $t")
            }
        }

        case Let(x, t, e1, e2) => {
            val t1 = typeOf(e1, alpha)
            if (typeEquals(t1, t)){
                val newAlpha = alpha + (x -> t)
                typeOf(e2, newAlpha)
            } else {
                throw TypeErrorException(s"Let binding has type $t whereas it is bound to expression of type $t1")
            }
        }

        case FunDef(x, t1, e) => {
            val newAlpha = alpha + (x -> t1)
            val t2 = typeOf(e, newAlpha)
            FunType(t1, t2)
        }

        case FunCall(e1, e2) => {
            val ftype = typeOf(e1, alpha)
            ftype match {
                case FunType(t1, t2) => {
                    val argType = typeOf(e2, alpha)
                    if (typeEquals(argType, t1)){
                        t2
                    } else {
                        throw TypeErrorException(s"Call to function with incompatible argument type. Expected $t1, obtained $argType")
                    }
                }
                case _ => { throw TypeErrorException(s"Call to function but with a non function type $ftype")}

            }
        }

        case NewRef(e) => {
            //??? // YOUR CODE HERE
            val rtype = typeOf(e, alpha)           //rType case matches for NewReference
            rtype match {
                case NumType => {
                    RefType(typeOf(e, alpha))
                }
                case BoolType => {
                    RefType(typeOf(e, alpha))
                }
                case FunType(t1,t2) => {
                    RefType(typeOf(e, alpha))
                }
                case RefType(t1) => {
                    RefType(typeOf(e, alpha))
                }
                case _ => { throw TypeErrorException(s"Expected Type")}
            }
            
        }
        
        case AssignRef(e1, e2) => {
            //??? // YOUR CODE HERE
            val rtype = typeOf(e1, alpha)      //rType Case matches for AssignRefereremce
            rtype match {
                case RefType(t1) => {
                    typeOf(e2, alpha)
                }
                case _ => { throw TypeErrorException(s"Call to AssignRef but with a non reference type $rtype")}
            }
        }
        
        
        case DeRef(e) => {
            //??? // YOUR CODE HERE
            val rtype = typeOf(e, alpha)          //rType case matches for DeReference
            rtype match {
                case RefType(t1) => {
                    t1
                }
                case _ => { throw TypeErrorException(s"Call to DeRef but with a non reference type $rtype")}
            }
        }
        
    }
}


def typeOfProgram(p: Program) = p match {           //Given
    case TopLevel(e) => {
            val t = typeOf(e, Map())
            println(s"Program type computed successfully as $t")
            t
    }
}


```

Output:
```
defined [32mfunction[39m [36mtypeOf[39m
defined [32mfunction[39m [36mtypeOfProgram[39m
```

<!-- grade_id: cell-e6bbf61c650e47bb, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST

/* 
let x : ref(num) = NewRef(10 ) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */

val p1 = Let("x", RefType(NumType), NewRef(Const(10)), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t1 = typeOfProgram(TopLevel(p1))
assert(t1 == NumType, "Test 1 failed: answer should be NumType")
passed(2)
//END TEST
```

Output:
```
Program type computed successfully as NumType

*** Tests Passed (2 points) ***
[36mp1[39m: [32mLet[39m = [33mLet[39m(
  [32m"x"[39m,
  [33mRefType[39m(NumType),
  [33mNewRef[39m([33mConst[39m([32m10.0[39m)),
  [33mLet[39m([32m"dummy"[39m, NumType, [33mAssignRef[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m30.0[39m)), [33mDeRef[39m([33mIdent[39m([32m"x"[39m)))
)
[36mt1[39m: [32mType[39m = NumType
```

<!-- grade_id: cell-e80a253dd4c598ac, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p2 = Let("x", RefType(NumType), NewRef(fdef), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t2 = try {
   typeOfProgram(TopLevel(p2))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(2)
//END TEST
```

Output:
```
*** Tests Passed (2 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))
[36mp2[39m: [32mLet[39m = [33mLet[39m(
  [32m"x"[39m,
  [33mRefType[39m(NumType),
  [33mNewRef[39m([33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))),
  [33mLet[39m([32m"dummy"[39m, NumType, [33mAssignRef[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m30.0[39m)), [33mDeRef[39m([33mIdent[39m([32m"x"[39m)))
)
[36mt2[39m: [32mAny[39m = [32m"OK -- caught a type error exception: Let binding has type RefType(NumType) whereas it is bound to expression of type RefType(FunType(NumType,NumType))"[39m
```

<!-- grade_id: cell-f6584c1da70c0ae8, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p3 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t3 = try {
   typeOfProgram(

TopLevel(p3))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(2)
//END TEST
```

<!-- grade_id: cell-9060540ef32281bd, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(NewRef(35), 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p4 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", NumType, AssignRef(NewRef(Const(35)), Const(30) ), DeRef(Ident("x"))) )
val t4 =  typeOfProgram(TopLevel(p4))
assert(t4 == FunType(NumType, NumType), "Test failed")
passed(2)
//END TEST
```

Output:
```
Program type computed successfully as FunType(NumType,NumType)

*** Tests Passed (2 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))
[36mp4[39m: [32mLet[39m = [33mLet[39m(
  [32m"x"[39m,
  [33mRefType[39m([33mFunType[39m(NumType, NumType)),
  [33mNewRef[39m([33mFunDef[39m([32m"z"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m10.0[39m)))),
  [33mLet[39m(
    [32m"dummy"[39m,
    NumType,
    [33mAssignRef[39m([33mNewRef[39m([33mConst[39m([32m35.0[39m)), [33mConst[39m([32m30.0[39m)),
    [33mDeRef[39m([33mIdent[39m([32m"x"[39m))
  )
)
[36mt4[39m: [32mType[39m = [33mFunType[39m(NumType, NumType)
```

<!-- grade_id: cell-b63b128c12f0247d, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST

val y = Ident("y")
val x = Ident("x")
val z = Ident("z")
val t1 = RefType(NumType)
val t2 = FunType(t1, NumType)
val t3 = RefType(t2)

val v1 = FunCall(DeRef(y), x)
val v2 = Let("z", RefType(BoolType), NewRef(Geq(DeRef(x), Const(25))), v1)
val fdef = FunDef("z", RefType(NumType), Plus(DeRef(z), DeRef(x)))
val v3 = Let("y", t3, NewRef(fdef), v2)
val v4 = Let("x", t1, NewRef(Const(10)), v3)

val p5 = TopLevel(v4)
assert(typeOfProgram(p5) == NumType, "Test failed")
passed(2)

//END TEST
```

Output:
```
Program type computed successfully as NumType

*** Tests Passed (2 points) ***
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36mz[39m: [32mIdent[39m = [33mIdent[39m([32m"z"[39m)
[36mt1[39m: [32mRefType[39m = [33mRefType[39m(NumType)
[36mt2[39m: [32mFunType[39m = [33mFunType[39m([33mRefType[39m(NumType), NumType)
[36mt3[39m: [32mRefType[39m = [33mRefType[39m([33mFunType[39m([33mRefType[39m(NumType), NumType))
[36mv1[39m: [32mFunCall[39m = [33mFunCall[39m([33mDeRef[39m([33mIdent[39m([32m"y"[39m)), [33mIdent[39m([32m"x"[39m))
[36mv2[39m: [32mLet[39m = [33mLet[39m(
  [32m"z"[39m,
  [33mRefType[39m(BoolType),
  [33mNewRef[39m([33mGeq[39m([33mDeRef[39m([33mIdent[39m([32m"x"[39m)), [33mConst[39m([32m25.0[39m))),
  [33mFunCall[39m([33mDeRef[39m([33mIdent[39m([32m"y"[39m)), [33mIdent[39m([32m"x"[39m))
)
[36mfdef[39m: [32mFunDef[39
```

## Problem 2 (20 points): Objects in Scala

### A (12 points)
Consider the class below for an employee record in an organization that has three fields `name`, `eid` (employee id) and `salary`.

1) As given, scala will complain that class employee cannot extend the abstract class People. Add the missing methods to enable the class Employee to properly extend People

2) Write a factory pattern that takes in a formatted string of the form "John Doe,29585,33110" consisting of a string and two numbers, to extract the name and salaries for each employee. You may want to use `scala.util.matching.Regex` [https://www.scala-lang.org/api/2.9.2/scala/util/matching/Regex.html]
If the input fails to match the pattern, an `IllegalArgumentException` must be thrown

Companion objects are covered in scala book Chapter 4.3. Also lookup here: https://alvinalexander.com/scala/factory-pattern-in-scala-design-patterns

<!-- grade_id: cell-f80a491b407a7ab7, points: 0, grade: false, solution: true -->
```scala
import scala.util.matching.Regex

abstract class People {
    /* I need the ability to convert every person to a string */
    def toString: String
    /* I need the ability to check if two people are the same */
    def equals (p: People): Boolean
}

/* 
  Note that for an employee to be equal to an object of type People, that object
  must be an employee with matching eid number 
*/

class Employee( val name: String, val eid: Int, var salary: Int) extends People {
    /////??? // YOUR CODE HERE
    override def toString = name 
    def equals(p: People) = {
        if (p.isInstanceOf[Employee]) {
            p.asInstanceOf[Employee].eid == eid
        }
        else{
            false
        }
    }
}


object Employee {
    def apply(formattedInput: String): Employee = {
        /* 
        Input format must be 
        (full name with possible spaces)[optional whitespaces],[optional whitespaces](employee ID number)[optional whitespaces],[optional whitespaces](employee salary)
        */
        //??? // YOUR CODE HERE
        val employee = """([a-zA-Z\s]+?)\s*,\s*(\d+)\s*,\s*(\d+)""".r     //possible values 
        formattedInput match {                     //case matches for employee
            case employee(name, eid, salary) => new Employee(name, eid.toInt, salary.toInt)
            case _ => throw new IllegalArgumentException
        }
    }
}



```

Output:
```
[32mimport [39m[36mscala.util.matching.Regex

[39m
defined [32mclass[39m [36mPeople[39m
defined [32mclass[39m [36mEmployee[39m
defined [32mobject[39m [36mEmployee[39m
```

<!-- grade_id: cell-e1ec3bf75d6f6de5, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
println(e1.toString)
assert(e1.name == "Jonathan Doe", "Failed")
assert(e1.eid == 12893, "Failed")
assert(e1.salary == 39813, "Failed")
assert(e1.equals(new Employee("whoknows", 12893, 39813)))
passed(4)
//END TEST
```

Output:
```
Jonathan Doe

*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Jonathan Doe
```

<!-- grade_id: cell-9d823f9842fd95a7, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
try {
    val e2 = Employee("Badly Formatted String, 29812, $3102")
    assert(false, "Failed, the string is badly formatted but your code is OK with it.")
} catch {
    case e => println(s"Expected behavior seen")
}
passed(4)
//END TEST
```

Output:
```
Expected behavior seen

*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-00d1907d97c7da69, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Janet Maria Catherine Doe    , 12894   , 121331")
println(e1.toString)
assert(e1.name == "Janet Maria Catherine Doe", "Failed")
assert(e1.eid == 12894, "Failed")
assert(e1.salary == 121331, "Failed")
assert(e1.equals(new Employee("random", 12894, 121332)))
passed(4)
//END TEST
```

Output:
```
Janet Maria Catherine Doe

*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Janet Maria Catherine Doe
```

### B (8 points)
We wish to add a new class of people to our system. These are "Contractor" with the information associated should involve their name (string), contract ID (Integer) and contract supervisor (Employee).

Write a new class named `Contractor` whose constructor should take in the name of the person (String), their ID number (Integer) and Supervisor (Employee) and extend from People. 

A contractor object is equal to an object p of type People, only if p is actually an instance of contractor, and if their names are the same and they have the same contract ID number.

<!-- grade_id: cell-97ab958367cfbd0a, points: 0, grade: false, solution: true -->
```scala
//??? // YOUR CODE HERE


class Contractor( val name: String, val id: Int, var supervisor: Employee) extends People {
    override def toString = name
    def equals(p: People) = {
        if (p.isInstanceOf[Contractor]) {
            p.asInstanceOf[Contractor].name == name && p.asInstanceOf[Contractor].id == id
        }
        else {
            false
        }
    }
}
```

Output:
```
defined [32mclass[39m [36mContractor[39m
```

<!-- grade_id: cell-82cc00e175a12efe, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val c1 = new Contractor("John Contract Worker", 12093, e1)
assert( !c1.equals(e1), "Failed")
assert( c1.equals(c1), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Jonathan Doe
[36mc1[39m: [32mContractor[39m = John Contract Worker
```

<!-- grade_id: cell-62fa2e0b66f58862, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val e2 = Employee("Jane Doe, 12894, 43814")
val c1 = new Contractor("John Contract Worker", 12093, e2)
val c2 = new Contractor("John Contract Worker", 12093, e1)
assert( c1.equals(c2), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mEmployee[39m = Jonathan Doe
[36me2[39m: [32mEmployee[39m = Jane Doe
[36mc1[39m: [32mContractor[39m = John Contract Worker
[36mc2[39m: [32mContractor[39m = John Contract Worker
```

## Problem 3 (10 points): Traits in Scala

For this problem, we have defined two traits: `NumberOfLegs` that helps us define how many legs a given animal has and `WarmBlooded` that applies to warm blooded animals.

```scala
abstract class Animal 

trait NumberOfLegs {
    val nLegs: Int
    def getNumberOfLegs: Int = nLegs
}

trait WarmBlooded extends Animal {
    val bodyTempMaintained: Double
    def getBodyTemp: Double = bodyTempMaintained 
}
```

Output:
```
defined [32mclass[39m [36mAnimal[39m
defined [32mtrait[39m [36mNumberOfLegs[39m
defined [32mtrait[39m [36mWarmBlooded[39m
```

### A (5 points)

Define a class `Human` that inherits from `Animal` and mixes in the traits `NumberOfLegs` with `nLegs = 2` and `WarmBlooded` with `bodyTempMaintained = 98`. The class `Human` must take in a parameter called `name` of type `String` and implement a `getName` method without any parameters.

<!-- grade_id: cell-e7a4a9275bd24e24, points: 0, grade: false, solution: true -->
```scala
//??? // YOUR CODE HERE
class Human(val name: String, val bodyTempMaintained: Double = 98, val nLegs: Int = 2) extends WarmBlooded with NumberOfLegs {    //Definition of Human class that follows inherits 'Animal' and mixes traits from 'NumberOfLegs
    def getName = name                                                                                                            ///////sets nLegs = 2,  WarmBlooded, BodtTemp = 98.
}
```

Output:
```
defined [32mclass[39m [36mHuman[39m
```

<!-- grade_id: cell-f9b7ae8902326372, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val t1 = new Human("Jane Smith")
assert(t1.getNumberOfLegs == 2, "Your human does not have two legs")
assert(t1.bodyTempMaintained == 98.0, "Your human does not maintain a body temp of 98")
assert(t1.getBodyTemp == 98.0, "Your human's getBodyTemp Function is not working")
assert(t1.getName == "Jane Smith", "Your human's name is not setting correctly")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mt1[39m: [32mHuman[39m = ammonite.$sess.cmd62$Helper$Human@49f3c6bb
```

### B (5 points)
Define a class named `Table` that mixes in the trait `NumberOfLegs` with `nLegs = 4`. 

_Note:_ `Table` cannot mixin the trait `WarmBlooded` (make sure you understand why that is as a quick check of concepts).

<!-- grade_id: cell-8145df690b6f6588, points: 0, grade: false, solution: true -->
```scala
//??? // YOUR CODE HERE
class Table(val nLegs: Int = 4) extends NumberOfLegs            /////////Mixes trait from Warmblooded and nlegs = 4
```

Output:
```
defined [32mclass[39m [36mTable[39m
```

<!-- grade_id: cell-d34b29b062b93894, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val tbl = new Table()
assert(tbl.getNumberOfLegs == 4, "A Table must have four legs")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mtbl[39m: [32mTable[39m = ammonite.$sess.cmd64$Helper$Table@536c52fe
```

```scala
////////////////////////////////////////////////////////////////////////////////////
///////////////////////////////////////////////////////////////////////////////////
//Jason Nguyen - Assignment 9 - Sriram Sankaranarayanan - During CoVid-19 Outbreak 2020//
////////////////////////////////////////////////////////////////////////////////////
///////////////////////////////////////////////////////////////////////////////////

```