# CSCI 3155: Assignment 2

Topics covered: recursion and inductive definitions.

__YOUR NAME HERE__

```scala
// TEST HELPER: DO NOT EDIT
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (15 points total)

### A (5 points) Recursive `cubeIt`

Here is a definition of a recursive function to compute $n^3$ for an integer $n$.

$$\mbox{cubeIt}(n) = \left\{ \begin{array}{ll} 
0 & \mbox{if}\ n = 0 \\
\mbox{cubeIt}(n-1) + 3n^2 - 3n + 1,\ & \mbox{if}\ n > 0 \\
\mbox{cubeIt}(n+1) - 3 n^2 - 3 n - 1 & \mbox{if}\ n < 0 \\
\end{array}\right.$$

Implement this recursion as a scala function `cubeIt`. Note (a) your code does not have to be tail recursive (see part B below), and (b) Be careful: `n^2` in scala is not the same as `n*n`. The `^` operator does XOR in scala.

<!-- grade_id: answer_1a, points: 0, grade: false, solution: true -->
```scala
def cubeIt(n: Int): Int = { 
    //BEGIN SOLUTION
    if (n == 0){
        0
    } else if (n > 0){
        cubeIt(n-1) + 3 *n*n - 3*n + 1
    } else {
        cubeIt(n+1) - 3 *n*n - 3*n - 1
    }
    //END SOLUTION
}
```

Output:
```
defined [32mfunction[39m [36mcubeIt[39m
```

<!-- grade_id: tests_1a, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
assert(cubeIt(5) == 125)
assert(cubeIt(-5) == -125)
assert(cubeIt(-3) == -27)
assert(cubeIt(0) == 0)
// JUST TEST from -100 to 100
def checkCubeItFun(x: Int): Boolean = cubeIt(x) == x*x*x
assert( (-100 to 100).forall(checkCubeItFun))
passed(5)
// END TESTS
```

Output:
```
*** Tests Passed (5 points) ***
defined [32mfunction[39m [36mcheckCubeItFun[39m
```

### B (10 points) _Tail_ Recursive `cubeIt`
Now, write a tail recursive version of the `cubeIt` function above, called `cubeItTail`. 
We have setup a helper function with an *accumulator* argument for you, to help you out. Do not remove the @tailrec annotation. It is part of the autograder: you can learn more about it's meaning here -- https://www.scala-exercises.org/scala_tutorial/tail_recursion

<!-- grade_id: answer_1b, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec

@tailrec
def cubeItHelper(n: Int, acc: Int): Int = { /* Implement cubeIt as a tail recursion */
    //BEGIN SOLUTION
    if (n == 0) { acc }
    else { 
        if (n < 0){
            cubeItHelper(n+1, acc - 3 *n*n - 3 *n -1)
        } else {
            cubeItHelper(n-1, acc + 3 * n*n - 3 * n + 1)
        }
    }
    //END SOLUTION
}

/*- 
This is the main function that the users will use. 
It should call the helper with the correct initial value of acc argument. 
Just replace the ??? with the correct initial value for the acc argument.
-*/

def cubeItTail(n: Int): Int = { 
    cubeItHelper(n, 
                //BEGIN SOLUTION
                 0
                //END SOLUTION
                )
}

```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec

[39m
defined [32mfunction[39m [36mcubeItHelper[39m
defined [32mfunction[39m [36mcubeItTail[39m
```

<!-- grade_id: tests_1b, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(cubeItTail(5) == 125)
assert(cubeItTail(-5) == -125)
assert(cubeItTail(-3) == -27)
assert(cubeItTail(0) == 0)
assert( (-100 to 100).forall( x => (cubeItTail(x) == x * x * x)), "Test Passed!")
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
```

## Problem 2 (10 points)

We have given scala code using a while loop that converts a bit string with 0s and 1s into its decimal equivalent.

```scala
def convertToDecimal(bStr : String): Int = {
    var bitStr = bStr // Make a copy of the string
    var acc = 0 // Initialize the accumulat to 0
    while (! bitStr.isEmpty) {
        if (bitStr(0) == '0'){ // Is the first character a 0?
            acc = 2 * acc // Update acc
            bitStr = bitStr.substring(1) // Remove the very first character from the string
        } else {
            acc = 2 * acc + 1 // Update acc
            bitStr = bitStr.substring(1)
        }
    }
    return acc
}
```

Output:
```
defined [32mfunction[39m [36mconvertToDecimal[39m
```

```scala
println(convertToDecimal("1001"))
println(convertToDecimal("101"))
println(convertToDecimal("0000"))
println(convertToDecimal("110110111010001"))
println(convertToDecimal("1000000001"))
```

Output:
```
9
5
0
28113
513
```

Note: (a) `string.substring(j)` selects the suffix of string from $j^{th}$ position (positions start at 0) to end; (b) note `string(0)` is the first character of the string;
and (c) `string.isEmpty` checks if a string is empty. 

Implement a *tail recursive* function called `dHelper` over "bit" strings (strings whose characters can only be 0 or 1) that converts them from binary to a decimal number using an accumulator argument. For your convenience, we have also translated the while loop into a mathematical definition that you can use.

$$\text{dHelper}(\text{bStr}, \text{acc}) = \left\{ \begin{array}{ll} 
\text{acc}, & \text{if bStr is empty}\\
\text{dHelper}(\text{bStr.substring}(1) , 2 *\text{acc} + 1 ), & \text{if}\ \text{bStr is not empty and  bStr(0) = '1'} \\
\text{dHelper}(\text{bStr.substring}(1) , 2 *\text{acc}  ), & \text{if bStr is not empty and bStr(0) = '0'} \\
\end{array}\right.$$



<!-- grade_id: answer_2, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec

@tailrec
def dHelper(bitStr: String, acc:Int): Int = {
    //BEGIN SOLUTION
    if (bitStr.isEmpty) { acc }
    else if (bitStr(0) == '1') {
        dHelper(bitStr.substring(1), 2 * acc + 1)
    } else {
        dHelper(bitStr.substring(1), 2 * acc)
    }
    //END SOLUTION
}

def toDecimal(bitStr: String) = dHelper(bitStr, 0)
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec

[39m
defined [32mfunction[39m [36mdHelper[39m
defined [32mfunction[39m [36mtoDecimal[39m
```

<!-- grade_id: test_2, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(toDecimal("0") == 0)
assert(toDecimal("10") == 2)
assert(toDecimal("1001") == 9)
assert((0 to 100).forall ( x => ( toDecimal(x.toBinaryString) == x)) )
//END TEST
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

## Problem 3 (20 points)

Convert the following inductive definition for regular expressions into a grammar first and then into a set of scala classes.

A regular expression is defined inductively as follows:
- Any string s is an "atomic" regular expression.
- If $r_1$, $r_2$ are regular expressions then so are
  - The concatenation $r_1 ; r_2$, 
  - The disjunction $r_1 | r_2$, and 
  - The conjunction $r_1 \& r_2$.
- If $r$ is a regular expression, then its Kleene star $r^*$ is also a regular expression.

Use the constructor symbols $\text{Atom}(s)$ to denote an atomic regular expression, $\text{Concat}(r_1, r_2)$ for the ";" operator, $\text{Or}(r_1, r_2)$ for the "|" operator, $\text{And}(r_1, r_2)$ for the "&" operator and $\text{Star}(r)$ for the Kleene-star operator. You may use
the nonterminal $\textbf{string}$ without definition to denote a string of characters.

| Operator | Constructor Symbol to Use |
| -------- | --------------------------|
| atomic string | Atom |
| ; (concatenation operator) | Concat | 
| \| (or operator) |  Or |
| \& (and operator) | And |
| \* (Kleene Star) | Star |




### A (7 point)
Write the grammar using constructor symbols for the inductive definition above. 

__Tip:__ you can examine the notes or the recitation notebooks to see how we typeset the grammar. There are no tests for this because it will be manualy graded.

<!-- grade_id: answer_3a, points: 7, grade: true, solution: true -->
$$\begin{array}{rcl}
\textbf{Regex} & \rightarrow & {Atom}(\textbf{string}) \\
& | & {Concat}(\textbf{Regex}, \textbf{Regex}) \\
& | & {Or}(\textbf{Regex}, \textbf{Regex}) \\
& | & {And}(\textbf{Regex}, \textbf{Regex}) \\
& | & {Star}(\textbf{Regex}) \\
\end{array}$$

### B (7 points)
Define the structure as a set of case class in scala.


<!-- grade_id: answer_3b, points: 0, grade: false, solution: true -->
```scala
sealed trait Regex
// Use constructors: Atom, Concat, Or, And and Star
//BEGIN SOLUTION
case class Atom(s: String) extends Regex
case class Concat(r1: Regex, r2: Regex) extends Regex
case class Or(r1: Regex, r2: Regex) extends Regex
case class And(r1: Regex, r2: Regex) extends Regex
case class Star(r: Regex) extends Regex
//END SOLUTION
```

Output:
```
defined [32mtrait[39m [36mRegex[39m
defined [32mclass[39m [36mAtom[39m
defined [32mclass[39m [36mConcat[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mStar[39m
```

<!-- grade_id: tests_3b, points: 7, grade: true, solution: false -->
```scala
// BEGIN TESTS
val v1 = Atom("Hello")
val v2 = Concat(v1, v1)
val v3 = Or(v1, v2)
val v4 = Star(v3)
val v5 = Or(v1, v4)
passed(7)
// END TESTS
```

Output:
```
*** Tests Passed (7 points) ***
[36mv1[39m: [32mAtom[39m = [33mAtom[39m([32m"Hello"[39m)
[36mv2[39m: [32mConcat[39m = [33mConcat[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m))
[36mv3[39m: [32mOr[39m = [33mOr[39m([33mAtom[39m([32m"Hello"[39m), [33mConcat[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m)))
[36mv4[39m: [32mStar[39m = [33mStar[39m([33mOr[39m([33mAtom[39m([32m"Hello"[39m), [33mConcat[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m))))
[36mv5[39m: [32mOr[39m = [33mOr[39m(
  [33mAtom[39m([32m"Hello"[39m),
  [33mStar[39m([33mOr[39m([33mAtom[39m([32m"Hello"[39m), [33mConcat[39m([33mAtom[39m([32m"Hello"[39m), [33mAtom[39m([32m"Hello"[39m))))
)
```

### C (6 points)
Write down the representation of the regular expression in Scala. Your cell must define a term that should be called `finalAnswerC`.

$$ (\text{"hello"})^* ;  ((\text{"scala"}; \text{"best"})^*) $$


<!-- grade_id: answer_3c, points: 0, grade: false, solution: true -->
```scala
val finalAnswerC = {
    // BEGIN SOLUTION
    val v1 = Atom("hello")
    val v2 = Concat(Atom("scala"), Atom("best"))
    val v3 = Star(v1)
    val v4 = Star(v2)
    val v5 = Concat(v3, v4)
    v5
    // END SOLUTION
}
```

Output:
```
[36mfinalAnswerC[39m: [32mConcat[39m = [33mConcat[39m(
  [33mStar[39m([33mAtom[39m([32m"hello"[39m)),
  [33mStar[39m([33mConcat[39m([33mAtom[39m([32m"scala"[39m), [33mAtom[39m([32m"best"[39m)))
)
```

<!-- grade_id: tests_3c, points: 6, grade: true, solution: false -->
```scala
// BEGIN TESTS
def munge(r: Regex): String = r match {
    case Atom(st) => "$A$"+st
    case Concat(r1, r2) => munge(r1)+"$C$"+munge(r2)
    case Star(r) => "$K$"+munge(r)+"$S$"
    case Or(r1, r2) => munge(r1)+"_O_"+munge(r2)
    case And(r1,r2) => munge(r1)+"/\\"+ munge(r2)
}

assert(munge(finalAnswerC) == "$K$$A$hello$S$$C$$K$$A$scala$C$$A$best$S$", "Test failed: you should seek help from the course staff to help you debug this problem, please")
passed(6)
// END TESTS
```

Output:
```
*** Tests Passed (6 points) ***
defined [32mfunction[39m [36mmunge[39m
```